```python
# Initialize Otter
import otter
grader = otter.Notebook("DSA2025Test.ipynb")
```

# DSA 2025 Summer School Admittance Check

Thanks for your interest in attending DSA 2025 Ibadan, Nigeria. To attend the summer school you have to have some level of basic Python proficiency. Completing the following notebook should ensure you have the right kind of background to benefit maximally from the Summer School. See you in Ibadan!

## Instructions
1. Complete each function according to the provided specifications
2. Run all cells to verify your solutions
3. All tests must pass to generate a submission
4. Save your work before submitting

```python
# Run these once ... To be on a safe side
!pip install nose
!pip install otter-grader
import IPython
from IPython import get_ipython
# Import the good stuff
import pandas as pd
import numpy as np
import math
from unittest import TestCase
assert_equal = TestCase().assertEqual
import otter
from collections import Counter

grader = otter.Notebook()
```

Output:
```
Requirement already satisfied: nose in c:\users\alvin\appdata\local\programs\python\python312\lib\site-packages (1.3.7)
Requirement already satisfied: otter-grader in c:\users\alvin\appdata\local\programs\python\python312\lib\site-packages (6.1.2)
Requirement already satisfied: click<9.0.0,>=8.1.7 in c:\users\alvin\appdata\local\programs\python\python312\lib\site-packages (from otter-grader) (8.1.8)
Requirement already satisfied: dill>=0.3.0 in c:\users\alvin\appdata\local\programs\python\python312\lib\site-packages (from otter-grader) (0.3.9)
Requirement already satisfied: fica>=0.4.1 in c:\users\alvin\appdata\local\programs\python\python312\lib\site-packages (from otter-grader) (0.4.1)
Requirement already satisfied: ipylab<2.0.0,>=1.0.0 in c:\users\alvin\appdata\local\programs\python\python312\lib\site-packages (from otter-grader) (1.0.0)
Requirement already satisfied: ipython in c:\users\alvin\appdata\local\programs\python\python312\lib\site-packages (from otter-grader) (9.0.1)
Requ
```

<!-- BEGIN QUESTION -->

**Question 1:** 
Write a Python function to return a tuple of even and odd numbers 

1.   List item
2.   List item

given an integer input.</br>
Example: For input 10, the output should be ([2, 4, 6, 8, 10], [1, 3, 5, 7, 9]).

```python
def even_odd_numbers(n):
    """
    Returns a tuple of even and odd numbers up to n, starting from 1.

    Args:
        n (int): The upper limit.

    Returns:
        tuple: A tuple of two lists: (evens, odds).
    """
    evens = [i for i in range(1, n + 1) if i % 2 == 0]
    odds = [i for i in range(1, n + 1) if i % 2 != 0]
    return evens, odds

#Testing the function
print(even_odd_numbers(10))
```

Output:
```
([2, 4, 6, 8, 10], [1, 3, 5, 7, 9])
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2: Create a dictionary to store the frequency of each word in the paragraph provided.**</br>
Example: For the paragraph "Data science is fun. Data science is useful.", the output should be {'Data': 2, 'science': 2, 'is': 2, 'fun': 1, 'useful': 1}.

```python
import string

def word_frequency(paragraph):
    """
    Returns a dictionary of word frequencies in a paragraph, treating words case-insensitively.

    Args:
        paragraph (str): The input paragraph.

    Returns:
        dict: A dictionary with words as keys and their frequencies as values.
    """
    #Converting the paragraph to lowercase forr case-insensitive counting
    paragraph = paragraph.lower()
    
    #Removes punctuations frosm the paragraphh
    translator = str.maketrans('', '', string.punctuation)
    cleaned_paragraph = paragraph.translate(translator)
    
    #Splits the cleaned paragraph into words
    words = cleaned_paragraph.split()
    
    #Counting the frequency of each word
    frequency = {}
    for word in words:
        frequency[word] = frequency.get(word, 0) + 1
        
    return frequency

#Example of usage
paragraph = "Data science is fun. Data science is useful."
print(word_frequency(paragraph))
```

Output:
```
{'data': 2, 'science': 2, 'is': 2, 'fun': 1, 'useful': 1}
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 3: Extract the values from the dictionary above into a list and sort them in descending order.**</br>
Example: For the dictionary {'a': 3, 'b': 1, 'c': 2}, the output should be [3, 2, 1]

```python
def extract_and_sort_values(dictionary):
    """
    Extracts and sorts dictionary values in descending order.

    Args:
        dictionary (dict): The input dictionary.

    Returns:
        list: A list of sorted values in descending order.
    """
    return sorted(dictionary.values(), reverse=True)

#Example usage
example_dict = {'a': 3, 'b': 1, 'c': 2}
sorted_values = extract_and_sort_values(example_dict)
print(sorted_values)
```

Output:
```
[3, 2, 1]
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 4: Merge the two lists (from Questions 3 and 4) into a dictionary where the keys are the sorted keys and the values are the sorted values.**</br>
Example: For lists ['a', 'b', 'c'] and [3, 2, 1], the output should be {'a': 3, 'b': 2, 'c': 1}

```python
def merge_lists_to_dict(keys, values):
    """
    Merges two lists into a dictionary with keys sorted in ascending order.

    Args:
        keys (list): List of keys.
        values (list): List of values.

    Returns:
        dict: A dictionary with keys and corresponding values.
    """
    #Sortingg the keys and get theirs corresponding indices
    sorted_key_indices = sorted(range(len(keys)), key=lambda i: keys[i])
    
    #Creating the merged dictionary using the sorted keys and original values
    merged_dict = {keys[i]: values[i] for i in sorted_key_indices}
    
    return merged_dict

#Example usage
keys = ['x', 'y']
values = [10, 20]
result = merge_lists_to_dict(keys, values)
print(result)
```

Output:
```
{'x': 10, 'y': 20}
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 5:Write a function least_common_multiple that takes two inputs a and b and returns the least common multiple of the two numbers.**</br>
Example: For input (4, 6), the output should be 12

```python
def least_common_multiple(a, b):
    """
    Returns the least common multiple of two numbers.

    Args:
        a (int): First number.
        b (int): Second number.

    Returns:
        int: The LCM of a and b.
    """
    return math.lcm(a, b)

#Example usage
lcm = least_common_multiple(4, 6)
print(lcm)
```

Output:
```
12
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 6:Write a function get_nearest_farthest that takes in a point of interest (x, y) and a list of points [(x1, y1), (x2, y2), ...] and returns the indices of the nearest and farthest points from the point of interest.** </br>
Example: For pt = (0, 0) and points = [(1, 1), (3, 3), (-1, -1)], the output should be (0, 1)

```python
def get_nearest_farthest(pt, points):
    """
    Returns the indices of the nearest and farthest points from a point of interest.

    Args:
        pt (tuple): The point of interest (x, y).
        points (list): List of points [(x1, y1), (x2, y2), ...].

    Returns:
        tuple: Indices of the nearest and farthest points.
    """
    #Convertss the lists of points to a NumPy array
    points_array = np.array(points)
    
    #Calculate the squared Euclidean distances from the point of interest to all other points
    squared_distances = np.sum((points_array - pt) ** 2, axis=1)
    
    #Identify the indices of the minimum and maximum squared distances
    nearest_index = np.argmin(squared_distances)
    
    farthest_index = np.argmax(squared_distances)
    return nearest_index, farthest_index

#Example usage
pt = (0, 0)
points = [(1, 1), (3, 3), (-1, -1)]
nearest, farthest = get_nearest_farthest(pt, points)
print(nearest, farthest)
```

Output:
```
0 1
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 7:Write a function filter_divisible to return a list of numbers between 0 and a number N that are not perfectly divisible by q.** </br>

> Hint: If N is negative, use N = 20. 

Example: For N = 10 and q = 2, the output should be [1, 3, 5, 7, 9]

```python
def filter_divisible(N, q):
    """
    Returns a list of numbers between 0 and N that are not divisible by q.

    Args:
        N (int): The upper limit.
        q (int): The divisor.

    Returns:
        list: A list of non-divisible numbers.
    """
    #If N is negative, set N to 20
    if N < 0:
        N = 20
        
    #Generate the list of numbers not divisible by q
    return [i for i in range(N + 1) if i % q != 0]

#Example usage
result = filter_divisible(10, 2)
print(result)
```

Output:
```
[1, 3, 5, 7, 9]
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 8: Write a function flatten_and_unique that takes in a list of lists and outputs a sorted list of unique elements from all sublists.**

```python
def flatten_and_unique(list_of_lists):
    """
    Flattens a list of lists and returns a sorted list of unique elements.

    Args:
        list_of_lists (list): A list of lists.

    Returns:
        list: A sorted list of unique elements.
    """
    #Flatten the list of lists and extract unique elements
    unique_elements = set()
    for sublist in list_of_lists:
        unique_elements.update(sublist)
        
    #Return the sorted list of unique elements
    return sorted(unique_elements)

#Example usage
nested_list = [['b', 'a', 'c'], ['a', 'd'], ['e', 'b']]
result = flatten_and_unique(nested_list)
print(result)
```

Output:
```
['a', 'b', 'c', 'd', 'e']
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

<!-- END QUESTION -->

**The Extra Mile!**


Download the dataset Nigeria Food Prices (9.9M) 

> https://data.humdata.org/dataset/wfp-food-prices-for-nigeria





<!-- BEGIN QUESTION -->

**Question 9: Create a new column date_new from the date column, converting it to a datetime format.**

```python
import pandas as pd

# Load the dataset
file_path = 'wfp_food_prices_nga.csv'  # Replace with the actual file path
df = pd.read_csv(file_path,skiprows=[1])
df = df.drop(index=2)
df = df.reset_index(drop=True)

# Display the first few rows of the dataset
print("First 5 rows of the dataset:")
print(df.head())


```

Output:
```
First 5 rows of the dataset:
         date   admin1 admin2        market  market_id  latitude  longitude  \
0  2002-01-15  Katsina  Jibia   Jibia (CBM)       1038     13.08       7.24   
1  2002-01-15  Katsina  Jibia   Jibia (CBM)       1038     13.08       7.24   
2  2002-01-15  Katsina  Jibia   Jibia (CBM)       1038     13.08       7.24   
3  2002-01-15  Katsina  Jibia   Jibia (CBM)       1038     13.08       7.24   
4  2002-01-15   Sokoto   Gada  Illela (CBM)       1037     13.64       5.28   

             category      commodity  commodity_id unit priceflag  pricetype  \
0  cereals and tubers          Maize            51   KG    actual  Wholesale   
1  cereals and tubers         Millet            73   KG    actual  Wholesale   
2  cereals and tubers        Sorghum            65   KG    actual  Wholesale   
3     pulses and nuts  Beans (niebe)           120   KG    actual  Wholesale   
4  cereals and tubers          Maize            51   KG    actual  Wholesale   

  currency   pr
```

```python
def create_date_new(df):
    """
    Creates a new column `date_new` from the `date` column, converting it to datetime format.

    Args:
        df (pd.DataFrame): The input DataFrame containing a 'date' column.

    Returns:
        pd.DataFrame: The DataFrame with the new `date_new` column.
    """
    #Convert the date column to datetime format and create the date_new column
    df['date_new'] = pd.to_datetime(df['date'], errors='coerce')
    return df

df.head()
```

Output:
```
date   admin1 admin2        market  market_id  latitude  longitude  \
0  2002-01-15  Katsina  Jibia   Jibia (CBM)       1038     13.08       7.24   
1  2002-01-15  Katsina  Jibia   Jibia (CBM)       1038     13.08       7.24   
2  2002-01-15  Katsina  Jibia   Jibia (CBM)       1038     13.08       7.24   
3  2002-01-15  Katsina  Jibia   Jibia (CBM)       1038     13.08       7.24   
4  2002-01-15   Sokoto   Gada  Illela (CBM)       1037     13.64       5.28   

             category      commodity  commodity_id unit priceflag  pricetype  \
0  cereals and tubers          Maize            51   KG    actual  Wholesale   
1  cereals and tubers         Millet            73   KG    actual  Wholesale   
2  cereals and tubers        Sorghum            65   KG    actual  Wholesale   
3     pulses and nuts  Beans (niebe)           120   KG    actual  Wholesale   
4  cereals and tubers          Maize            51   KG    actual  Wholesale   

  currency   price  usdprice   date_new  
0      NGN 
```

```python
grader.check("q9")
```

Output:
```
q9 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 10:Split the dataframe into two separate dataframes based on the category column: one for cereals and tubers and another for pulses and nuts**

```python
print(df['category'].unique())
```

Output:
```
['cereals and tubers' 'pulses and nuts' 'non-food' 'oil and fats'
 'meat, fish and eggs' 'milk and dairy' 'vegetables and fruits'
 'miscellaneous food']
```

```python
def split_dataframes(df):
    """
    Splits the dataframe into separate DataFrames for 'cereals and tubers' and 'pulses and nuts'.

    Args:
        df (pd.DataFrame): The input DataFrame.

    Returns:
        tuple: A tuple containing two DataFrames: (cereals_df, pulses_df).
    """
    cereals_df = df[df['category'] == 'cereals and tubers']
    pulses_df = df[df['category'] == 'pulses and nuts']
    return cereals_df, pulses_df
```

```python
grader.check("q10")
```

Output:
```
q10 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 11:Calculate the mean, median, and mode of the price and usdprice columns for each category**

```python
def calculate_price_stats(df):
    """
    Calculates the mean, median, and mode of the `price` and `usdprice` columns for each `category`.

    Args:
        df (pd.DataFrame): The input dataframe. It must contain the following columns:
            - 'category'
            - 'price'
            - 'usdprice'

    Returns:
        tuple: A tuple of two DataFrames: (price_stats, usdprice_stats)
            - price_stats includes ['category', 'mean_price', 'median_price', 'mode_price']
            - usdprice_stats includes ['category', 'mean_usdprice', 'median_usdprice', 'mode_usdprice']
    """
    
    #Define a helper function to compute the mode.
    def get_mode(series):
        m = series.mode()
        return m.iloc[0] if not m.empty else None

    #Compute the statistics for price   
    price_stats = df.groupby('category', as_index=False).agg(
        mean_price=('price', 'mean'),
        median_price=('price', 'median'),
        mode_price=('price', get_mode)
    )
    
    #Compute tthe statistics for usdprice
    usdprice_stats = df.groupby('category', as_index=False).agg(
        mean_usdprice=('usdprice', 'mean'),
        median_usdprice=('usdprice', 'median'),
        mode_usdprice=('usdprice', get_mode)
    )
    
    #Assign the results to global variables
    globals()['price_stats'] = price_stats
    globals()['usdprice_stats'] = usdprice_stats
    
    #Return the tupless
    return price_stats, usdprice_stats
```

```python
grader.check("q11")
```

Output:
```
q11 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 12:Merge the two dataframes (from Question 11) back into one dataframe and save it as `merged_data.csv`.**

```python
def merge_dataframes(cereals_df, pulses_df):
    """
    Merges the two dataframes back into one dataframe and saves it as `merged_data.csv`.

    Args:
        cereals_df (pd.DataFrame): The cereals and tubers dataframe.
        pulses_df (pd.DataFrame): The pulses and nuts dataframe.

    Returns:
        pd.DataFrame: The merged dataframe.
    """
    #Merge the dataframes by concatenating them along rows
    merged_df = pd.concat([cereals_df, pulses_df], ignore_index=True)
    
    #Save the merged dataframe as a CSV file
    merged_df.to_csv("merged_data.csv", index=False)
    
    #Assign merged_df to the global variable to exists globally
    globals()['merged_df'] = merged_df
    
    return merged_df
```

```python
grader.check("q12")
```

Output:
```
q12 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 13:Open the `merged_data.csv` file and select only the `date_new`, `market`, `commodity`, `price`, and `usdprice` columns**

```python
def select_columns():
    """
    Opens the `merged_data.csv` file and selects only the `date_new`, `market`, `commodity`, `price`, and `usdprice` columns.

    Returns:
        pd.DataFrame: The dataframe with selected columns.
    """
    #Reading the merged CSV file
    df = pd.read_csv("merged_data.csv")
    
    #Selecting the desired columns
    selected_columns = df[["date_new", "market", "commodity", "price", "usdprice"]]
    
    #Assigningg the DataFrame to a global variable
    globals()['selected_columns'] = selected_columns
    
    return selected_columns
```

```python
grader.check("q13")
```

Output:
```
q13 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 14:Group the data by `admin1` (state) and calculate the average `price` and `usdprice` for each state**

```python
def calculate_state_avg_prices(df):
    """
    Groups the data by `admin1` (state) and calculates the average `price` and `usdprice` for each state.

    Args:
        df (pd.DataFrame): The input dataframe.

    Returns:
        pd.DataFrame: The dataframe with average prices by state.
    """
    state_avg_prices = df.groupby('admin1', as_index=False)[['price', 'usdprice']].mean()
    
    #Assigning to global variable
    globals()['state_avg_prices'] = state_avg_prices
    return state_avg_prices
```

```python
grader.check("q14")
```

Output:
```
q14 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 15:Identify the top 5 markets with the highest average `usdprice` for `Rice (imported)`.**

```python
def identify_top_markets(df):
    """
    Identifies the top 5 markets with the highest average `usdprice` for `Rice (imported)`.

    Args:
        df (pd.DataFrame): The input dataframe.

    Returns:
        pd.Series: A series with the top 5 markets and their average `usdprice`.
    """
    #Filters the dataframe to only include rows where commodity is Rice that is imported
    rice_df = df[df["commodity"] == "Rice (imported)"]
    
    #Groups using the market and compute the average usdprice
    market_avg = rice_df.groupby("market", as_index=False)["usdprice"].mean()
    
    #Sorting with usdprice in descending order and take the top 5 markets
    top_markets_df = market_avg.sort_values(by="usdprice", ascending=False).head(5)
    
    #Converting thes results into a series with market as indexs and average usdprice as values
    top_markets = top_markets_df.set_index("market")["usdprice"]
    
    #Assign the series to a global variable
    globals()['top_markets'] = top_markets
    
    return top_markets
```

```python
grader.check("q15")
```

Output:
```
q15 results: All test cases passed!
```

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1 results: All test cases passed!

q10 results: All test cases passed!

q11 results: All test cases passed!

q12 results: All test cases passed!

q13 results: All test cases passed!

q14 results: All test cases passed!

q15 results: All test cases passed!

q2 results: All test cases passed!

q3 results: All test cases passed!

q4 results: All test cases passed!

q5 results: All test cases passed!

q6 results: All test cases passed!

q7 results: All test cases passed!

q8 results: All test cases passed!

q9 results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Complete each function according to the provided specifications. Make sure all tests pass.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Output:
```
C:\Users\Alvin\AppData\Local\Programs\Python\Python312\Lib\site-packages\otter\check\notebook.py:494: UserWarning: Could not locate a PDF to include
  warnings.warn("Could not locate a PDF to include")
VBox(children=(HTML(value='<p style="margin: 0">A PDF of your notebook could not be generated. Please acknowle…
```

Tests for q1 (points: 1)
```
>>> assert even_odd_numbers(10) == ([2, 4, 6, 8, 10], [1, 3, 5, 7, 9])
```
```
>>> assert even_odd_numbers(5) == ([2, 4], [1, 3, 5])
```


Tests for q10 (points: 1)
```
>>> (cereals_df, pulses_df) = split_dataframes(df)
>>> assert len(cereals_df) > 0, 'Cereals and tubers dataframe is empty'
```
```
>>> assert len(pulses_df) > 0, 'Pulses and nuts dataframe is empty'
```


Tests for q11 (points: 1)
```
>>> (price_stats, usdprice_stats) = calculate_price_stats(df)
>>> assert not price_stats.empty, 'Price statistics dataframe is empty'
```
```
>>> assert not usdprice_stats.empty, 'USD price statistics dataframe is empty'
```


Tests for q12 (points: 1)
```
>>> merged_df = merge_dataframes(cereals_df, pulses_df)
>>> assert not merged_df.empty, 'Merged dataframe is empty'
```
```
>>> assert len(merged_df) == len(cereals_df) + len(pulses_df), 'Merged dataframe has incorrect number of rows'
```


Tests for q13 (points: 1)
```
>>> selected_columns = select_columns()
>>> assert not selected_columns.empty, 'Selected columns dataframe is empty'
```
```
>>> assert list(selected_columns.columns) == ['date_new', 'market', 'commodity', 'price', 'usdprice'], 'Incorrect columns selected'
```


Tests for q14 (points: 1)
```
>>> state_avg_prices = calculate_state_avg_prices(df)
>>> assert not state_avg_prices.empty, 'State average prices dataframe is empty'
```
```
>>> assert 'price' in state_avg_prices.columns, 'Price column not found in state average prices'
```


Tests for q15 (points: 1)
```
>>> top_markets = identify_top_markets(df)
>>> assert not top_markets.empty, 'Top markets dataframe is empty'
```
```
>>> assert len(top_markets) == 5, 'Incorrect number of top markets'
```


Tests for q2 (points: 1)
```
>>> assert word_frequency('Data science is fun. Data science is useful.') == {'data': 2, 'science': 2, 'is': 2, 'fun': 1, 'useful': 1}
```
```
>>> assert word_frequency('Hello hello world') == {'hello': 2, 'world': 1}
```


Tests for q3 (points: 1)
```
>>> assert extract_and_sort_values({'a': 3, 'b': 1, 'c': 2}) == [3, 2, 1]
```
```
>>> assert extract_and_sort_values({'x': 10, 'y': 5, 'z': 20}) == [20, 10, 5]
```


Tests for q4 (points: 1)
```
>>> assert merge_lists_to_dict(['a', 'b', 'c'], [3, 2, 1]) == {'a': 3, 'b': 2, 'c': 1}
```
```
>>> assert merge_lists_to_dict(['x', 'y'], [10, 20]) == {'x': 10, 'y': 20}
```


Tests for q5 (points: 1)
```
>>> assert least_common_multiple(4, 6) == 12
```
```
>>> assert least_common_multiple(5, 10) == 10
```


Tests for q6 (points: 1)
```
>>> assert get_nearest_farthest((0, 0), [(1, 1), (3, 3), (-1, -1)]) == (0, 1)
```
```
>>> assert get_nearest_farthest((0, 0), [(2, 2), (1, 1), (3, 3)]) == (1, 2)
```


Tests for q7 (points: 1)
```
>>> assert filter_divisible(10, 2) == [1, 3, 5, 7, 9]
```
```
>>> assert filter_divisible(15, 3) == [1, 2, 4, 5, 7, 8, 10, 11, 13, 14]
```


Tests for q8 (points: 1)
```
>>> assert flatten_and_unique([[2, 13, 44], [6, 7, 2]]) == [2, 6, 7, 13, 44]
```
```
>>> assert flatten_and_unique([[1, 2, 3], [3, 4, 5]]) == [1, 2, 3, 4, 5]
```


Tests for q9 (points: 1)
```
>>> df = create_date_new(df)
>>> assert 'date_new' in df.columns, "Column 'date_new' not found in the dataframe"
```
```
>>> assert df['date_new'].dtype == 'datetime64[ns]', "Column 'date_new' is not in datetime format"
```
